# *Radix Sort*

**Prerequisites**: [*Basics of Loops, Lists & Functions & Counting Sort algorithm*]

## Exercise 1

*Before moving onto the Radix Sort algo, let's first implement counting sort which will be a subroutine.*

*Complete the following count() function that will return an array that contains the count values of the elements in the input array. (place values are one's, ten's, hundred's and so on.)*

#### Code Cell
Below is the code cell for **Exercise 1**. Type in your code in the cell. 

<!-- grade_id: cell-2ec62f126b24dfbd, points: 0, grade: false, solution: true -->
```python
def count(array, place):
    
    count = [0] * 10
    size = len(array)
    
    ### BEGIN SOLUTION
    for i in range(0, size):
        index = array[i] // place
        count[index % 10] += 1
    ### END SOLUTION
    
    return count
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 1**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-933483ef34527a9c, points: 0, grade: true, solution: false -->
```python
assert count([12,14,22,15,20,30],1) == [2, 0, 2, 0, 1, 1, 0, 0, 0, 0]
assert count([10,20,30,11,22,33,21,22,23,100,200,300],10) == [3, 2, 5, 2, 0, 0, 0, 0, 0, 0]

print("All Passed")
###################################################
```

Output:
```
All Passed
```

----
## Exercise 2


*Now let's calculate the cumulative count.*

*Complete the following cumulative() function that will return an array that will have the cumulative count of elements in the input array.*

#### Code Cell
Below is the code cell for **Exercise 2**. Type your code in the cell.

<!-- grade_id: cell-95b9b126dd4a4a44, points: 0, grade: false, solution: true -->
```python
def cumulative(count):
    # Calculate cummulative count
    for i in range(1, 10):
        
        ### BEGIN SOLUTION
        count[i] += count[i - 1]
        ### END SOLUTION
    
    return count
```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 2**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-f502e8909dcf6eb8, points: 0, grade: true, solution: false -->
```python
assert cumulative(count([12,14,22,15,20,30],1)) == [2, 2, 4, 4, 5, 6, 6, 6, 6, 6]
assert cumulative(count([11,22,33,44,21,31,32,33],10)) == [0, 1, 3, 7, 8, 8, 8, 8, 8, 8]

print("All Passed")
###################################################
```

Output:
```
All Passed
```

----
## Exercise 3


*Since we have created two major functions of counting sort, now let's combine all, to implement Counting Sort algorithm.*

*Complete the following countingSort() function which will return a sorted version of a list based on a particular place value passed.*

#### Code Cell
Below is the code cell for **Exercise 3**. Type your code in the cell.

<!-- grade_id: cell-70a9b90a2f9fd107, points: 0, grade: false, solution: true -->
```python
def countingSort(array, place):

    size = len(array)
    output = [0] * size

    # Calculate count of elements
    counter = count(array,place) ## call to the count func. defined in Exercise 1
    
    # Calculate cummulative count
    cumulative(counter) ## call to the cumulative func. defined above in Exercise 2
    
    # Place the elements in sorted order
    i = size - 1
    while i >= 0:
        index = array[i] // place ## determine the index position on the place value basis
        ### BEGIN SOLUTION
        output[counter[index % 10] - 1] = array[i]
        counter[index % 10] -= 1
        ### END SOLUTION
        i -= 1 ## decrement the array index
        
    # store values of the output back into the original array
    for i in range(0, size):
        array[i] = output[i]
        
    return array ## this is the sorted version

```

#### Test Cases 
Run the below test cell to validate your code for **Exercise 3**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-33d4a01fd430ba7b, points: 0, grade: true, solution: false -->
```python
assert countingSort([10,20,22,33,19,190,1],1) == [10, 20, 190, 1, 22, 33, 19]
assert countingSort([10,20,30,40,50],10) == [10,20,30,40,50] 
assert countingSort([10,20,22,33,19,190,1],10) == [1, 10, 19, 20, 22, 33, 190]

print("All Passed")
###################################################
```

Output:
```
All Passed
```

----
## Exercise 4

*Since we have defined the subroutine of radix sort that is counting sort,*
*Now let's Understand the functioning of Radix Sort.*

*Complete the following radixSort() function to return a sorted version of the input array.*

#### Code Cell
Below is the code cell for **Exercise 4**. Type your code in the cell.

<!-- grade_id: cell-5d0ea47ca640c947, points: 0, grade: false, solution: true -->
```python
# We will be using the same countingSort() defined in exercise 3 to sort the elements in the basis of significant places

def radixSort(array):
    # Get maximum element
    max_element = max(array)
    
    # Apply counting sort to sort elements based on place value.
    place = 1
    ### BEGIN SOLUTION
    while max_element // place > 0:
        countingSort(array, place)
        place *= 10
    ### END SOLUTION
    return array
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 4**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-833f0ae200c2a3ea, points: 0, grade: true, solution: false -->
```python
assert radixSort([121, 432, 564, 23, 1, 45, 788]) == [1, 23, 45, 121, 432, 564, 788]
assert radixSort([100,200,1,552,255]) == [1, 100, 200, 255, 552]

print("All Passed!")
###################################################
```

Output:
```
All Passed!
```

----
## Exercise 5

*Fill in the missing pieces of code.*

#### Code Cell
Below is the code cell for **Exercise 5**. Type your code in the cell.

<!-- grade_id: cell-658b3178f675dcb0, points: 0, grade: false, solution: true -->
```python
def radix_Sort(arr): 
  
    # Find the maximum number to know number of digits 
    max1 = max(arr) 
  
    # perform counting sort on all the digits
    # Note that instead 
    # of passing digit number, exp is passed. exp is 10^i 
    # where i is current digit number 
    exp = 1

    ### BEGIN SOLUTION
    while max1/exp > 0: 
    ### END SOLUTION
     while ____________:
        countingSort(arr,exp) 
    ### BEGIN SOLUTION
        exp *= 10
    ### END SOLUTION
         _________
    return arr
```

#### Test Cases
Run the below test cell to validate your code for **Exercise 5**. The cell has test cases your code should pass. Please do not modify the contents of this cell.

<!-- grade_id: cell-6344067b20760ae0, points: 0, grade: true, solution: false -->
```python
assert radix_Sort([10,9,8,7,6,5,4]) == [4, 5, 6, 7, 8, 9, 10]
assert radix_Sort([100,200,1,552,255]) == [1, 100, 200, 255, 552]

print("All Passed!")
###################################################
```

Output:
```
All Passed!
```